<a href="https://colab.research.google.com/github/Shikhar-ss/Lab-Coursework-CE880/blob/master/problem1_week2.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

---
# Welcome to CE880
### This is your week-2 : Problem 1 notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 


---
## Question 1:

Write a function to pass a list of numbers and return log2 of the array in the list format. 
`Hint`: take input as list and inside the function covert it into numpy array then perform the function. `Example` : input `[1, 4, 8]` and and output should be:  `[0.0, 2.0, 3.0]` 

<!-- grade_id: myLog_ans, points: 0, grade: false, solution: true -->
```python
import numpy as np
from math import log2

def myLog(mylist):
    result = list(map(lambda x : log2(x),mylist))
    # np.array(result)
    return result
    raise NotImplementedError()
```

```python
in_array = [1, 3, 5, 10**8]
myLog(in_array)
```

Output:
```
[0.0, 1.584962500721156, 2.321928094887362, 26.575424759098897]
```

<!-- grade_id: myLog_test, points: 1, grade: true, solution: false -->
```python
"""Check your output"""
# Test case 1
my = [1, 3, 5, 10**8]
assert myLog(my) == [0.0, 1.584962500721156, 2.321928094887362, 26.575424759098897]

# Test case 2
my = [1, 100, 1000]
assert myLog(my) == [0.0, 6.643856189774724, 9.965784284662087]
```

---

## Question 2:

The Fibonacci numbers are the numbers in the following integer sequence.
`[0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, ...]`
In mathematical terms, the sequence $F_n$ of Fibonacci numbers is defined by the recurrence relation 

$$
F_n = F_{n-1} + F_{n-2}
$$
with seed values  
$$
 F_0 = 0 \& F_1 = 1.
$$


Your task here is to write a code for printing $n^{th}$ Fibonacci number
`Example`: Input: `myfibonacci(9)` and Output should be `34`. Remember your index starts from 0.

<!-- grade_id: myfibonacci_ans, points: 0, grade: false, solution: true -->
```python
def myfibonacci(n):
    """Function for nth Fibonacci number"""
    # YOUR CODE HERE
    if n == 0: return 0
    elif(n ==1) : return 1
    else:
      f,s= 0,1
      for i in range(n):
        f,s = s, f+s
    return f
      
    raise NotImplementedError()
```

```python
myfibonacci(3)
```

Output:
```
2
```

<!-- grade_id: myfibonacci_test, points: 1, grade: true, solution: false -->
```python
"""Check that correct fibonacci number"""
assert myfibonacci(0) == 0
assert myfibonacci(1) == 1
assert myfibonacci(2) == 1
assert myfibonacci(3) == 2
assert myfibonacci(4) == 3
assert myfibonacci(5) == 5
assert myfibonacci(6) == 8
assert myfibonacci(7) == 13
assert myfibonacci(8) == 21
assert myfibonacci(9) == 34
```

---
## Question 3: 

 Write a program, which accept a list with `nan`s in it and count the maximum number of consecutive repeated `nan`.
 Example: In the `a = [nan, nan, 0.1, nan, 0.16, 1, 0.16, 0.9999, 0.0001, 0.16, 0.101, nan, 0.16]` and answer is 2.

<!-- grade_id: max_nan_ans, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE


from math import isnan
from cmath import nan

def max_rep_nan(a):
    if len(a) == 0:
        return 0
    else:
        count_store = []
        count = 0
        
        for i in a:
            if isnan(i) == True:
               count += 1
            else:
                count_store.append(count)
                count = 0

        return max(count_store)

```

```python
from cmath import nan
a = [nan, nan, 0.1, nan, 0.16, 1, 0.16, 0.9999, 0.0001, 0.16, 0.101, nan, 0.16]
max_rep_nan(a)
```

Output:
```
2
```

<!-- grade_id: max_rep_test, points: 0, grade: true, solution: false -->
```python
"""Check your answer by running this cell"""
from cmath import nan
a = [nan, nan, nan, 0.1, nan, 0.16, 1, 0.16, 0.101, nan]
max_rep_nan(a) == 3

a = [nan, nan, nan, nan, nan, 0.16, 1, 0.16, 0.101, nan]
assert max_rep_nan(a) == 5

a = [nan, 10, nan, 0.1, nan, 0.16, 1, 0.16, 0.101, nan]
assert max_rep_nan(a) == 1

a = []
assert max_rep_nan(a) == 0
```

## Pandas Data Frame 


```python
from sklearn import datasets
import pandas as pd
```

```python
# import some data to play with
iris = datasets.load_iris()
x = pd.DataFrame(iris.data[:, :])  # we are taking all features.
y = iris.target
```

---
## Question 4:

Write a function to find the max number in the each column in the IRIS dataset and return the output as a list,

<!-- grade_id: filter_1_ans, points: 0, grade: false, solution: true -->
```python
def my_max_min(x):
    max_list = list(map(lambda i: max(x[i]), x.columns.values))
    return max_list
    raise NotImplementedError()
```

<!-- grade_id: filter_1_test, points: 1, grade: true, solution: false -->
```python
"""Check your output"""
assert my_max_min(x) == [7.9, 4.4, 6.9, 2.5]
```

---
## Question 5:


Write a function to calculate the `mean` and `standard deviation` of IRIS dataset (last column only). `Hint`: You can use NumPy aggregate functions. Create a list and append value to it and return the output. 

<!-- grade_id: stats_ans, points: 0, grade: false, solution: true -->
```python
def stats(x):
  vals = list(x[x.columns.values[-1]].agg([np.mean, np.std]))
  return vals
  raise NotImplementedError()
```

<!-- grade_id: stats_test, points: 2, grade: true, solution: false -->
```python
"""Check your output"""
# Test case
assert stats(x) == [1.1993333333333336, 0.7622376689603465]
```

```python

```